## Lab02: Logic

- Full Name:

- Student ID:

<font color='red'>**Some notes:**</font>

You can discuss ideas with classmates and find information from the internet, books, etc., but this homework must be your own. **Any detected cheating will result in a score of 0 for the entire course**

**How to submit your homework**

Before submitting, rerun the notebook (`Kernel` ->` Restart & Run All`).

Create folder named `Student_ID` (for example, if your student id is 1234567, you will name the folder as `1234567`) with the following structure:

   - Folder `Student_ID`
        - File `Student_ID.ipynb` 
        - Folder `models`, `Input`, `Output`and file `logic.py`
        - Folder `PS4` that include your works on exercise 4:
            - `source_code.py`: Source code (file .py)
            - `report.pdf`: Report (file .pdf)

Compress `Student_ID` folder with <font color='red'>**zip**</font> format.

<font color=red>Please, follow exactly all the requirements and instructions above.</font>



## How to do your homework: 

For questions 1-3, you will work directly in this notebook. The comment `# YOUR CODE HERE` will indicate the parts you need to complete (remember to comment out the line raise `NotImplementedError()` before submitting). Mathematical symbols must be presented in code according to the following symbol list:"

<center>
    
|  | Math symbols | Code | 
|:---|:----:|---|
| constant symbol | hcmus | ```Constant('hcmus')``` (must be lowercase) | 
| variable symbol | $x$ | ```Variable('$x')``` ( must be lowercase)  | 
| atomic formula | Rain, LocatedIn(hcmus, $x$) | ```Atom('Rain')```, ```Atom('LocatedIn', 'hcmus', '$x')``` <br> (predicate must be uppercase, arguments are symbols)|
| negation | $\neg$ Rain | ```Not(Atom('Rain'))```|
| conjunction| Rain $\land$ Snow | ```And(Atom('Rain'), Atom('Snow'))```|
| disjunction | Rain $\lor$ Snow | ```Or(Atom('Rain'), Atom('Snow'))```|
| implication | Rain $\Rightarrow$ Wet | ```Implies(Atom('Rain'), Atom('Wet'))```|
| equivalence | Rain $\Leftrightarrow$ Wet | ```Equiv(Atom('Rain'), Atom('Wet'))```|
| existential quantification | $\exists x$ LocatedIn(hcmus,$x$) | ```Exists('$x', Atom('LocatedIn', 'hcmus', '$x'))```|
| universal quantification | $\forall x$ MadeOfAtoms($x$) | ```Forall('$x', Atom('MadeOfAtoms', '$x'))```|    
| conjunction with many argument <br>(similar with disjunction)| A $\land$ B $\land$ C | ```AndList([Atom('A'), Atom('B'), Atom('C')]) ```|
  
    
</center>

```python
import pickle, gzip, os, random
import sys
import os
import collections

from logic import *
from typing import Tuple, List
```

```python
# import other modules if you need them 

```

## 0. Logic Testing Fucntion

```python

def checkFormula(name, predForm, preconditionForm=None):
    """
    Args:
        name: Name of this formula (used to load the models).
        predForm: The formula predicted in the submission.
        preconditionForm: Only consider models such that `preconditionForm` is true.
    """
    filename = os.path.join('models', name + '.pklz')
    objects, targetModels = pickle.load(gzip.open(filename)) 
    
    # If preconditionion exists, change the formula to
    preconditionPredForm = And(preconditionForm, predForm) if preconditionForm else predForm
    predModels = performModelChecking([preconditionPredForm], findAll=True, objects=objects)
    ok = True
    
    def hashkey(model): return tuple(sorted(str(atom) for atom in model))
    targetModelSet = set(hashkey(model) for model in targetModels)
    predModelSet = set(hashkey(model) for model in predModels)
    for model in targetModels:
        if hashkey(model) not in predModelSet:
            print("Your formula (%s) says the following model is FALSE, but it should be TRUE:" % predForm)
            ok = False
            printModel(model)
            return
    for model in predModels:
        if hashkey(model) not in targetModelSet:
            print("Your formula (%s) says the following model is FALSE, but it should be TRUE:" % predForm)
            return
        
    print(f"You matched the {len(targetModels)} models")
    print(f"Example model: {rstr(random.choice(targetModels))}")
    return True


```

## 1. Prepositional Logic (0.75 points)
In this session, you will practice transforming a natural languages into a propositional logic formula. For each successful assertion, you will receive 0.25 points.

For example: "If it 's rain, the road will be wet":

```
def rainWet():
    Rain = Atom('Rain') 
    Wet = Atom('Wet')
    return Implies(Rain, Wet)
```


##### **1a) If it's summer in California, it will not rain**

<!-- grade_id: TODO_1a, points: 0, grade: false, solution: true -->
```python
def formula1a() -> Formula:
    # Predicates to use:
    Summer = Atom('Summer')               # whether it's summer
    California = Atom('California')       # whether we're in California
    Rain = Atom('Rain')                   # whether it's raining
    
    # YOUR CODE HERE
    # raise NotImplementedError()
    return Implies(And(Summer, California), Not(Rain))
```

<!-- grade_id: GRADE_1a, points: 0.25, grade: true, solution: false -->
```python
assert checkFormula('1a', formula1a()) == True
```

Output:
```
You matched the 7 models
Example model: {'Summer'}
```

##### **1b) The road is wet if and only if the weather is rainy or the sprinkler is opened**

```python
def formula1b() -> Formula:
    # Predicates to use:
    Rain = Atom('Rain')              # whether it is raining
    Wet = Atom('Wet')                # whether it it wet
    Sprinklers = Atom('Sprinklers')  # whether the sprinklers are on
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
assert checkFormula('1b', formula1b()) == True
```

Output:
```
Cell In[7], line 1
----> 1 assert checkFormula('1b', formula1b()) == True

Cell In[6], line 7, in formula1b()
      5 Sprinklers = Atom('Sprinklers')  # whether the sprinklers are on
      6 # YOUR CODE HERE
----> 7 raise NotImplementedError()

NotImplementedError:
```

##### **1c) Day or night**

<!-- grade_id: TODO_1b, points: 0, grade: false, solution: true -->
```python
def formula1c() -> Formula:
    # Predicates to use:
    Day = Atom('Day')     # wh\\\\ether it's day
    Night = Atom('Night') # whether it's night
    # YOUR CODE HERE
    # raise NotImplementedError()
    return Or(And(Day, Not(Night)), And(Not(Day), Night))
```

<!-- grade_id: GRADE_1b, points: 0.25, grade: true, solution: false -->
```python
assert checkFormula('1c', formula1c()) == True
```

Output:
```
File c:\Python312\Lib\gzip.py:61, in open(filename, mode, compresslevel, encoding, errors, newline)
     59 gz_mode = mode.replace("t", "")
     60 if isinstance(filename, (str, bytes, os.PathLike)):
---> 61     binary_file = GzipFile(filename, gz_mode, compresslevel)
     62 elif hasattr(filename, "read") or hasattr(filename, "write"):
     63     binary_file = GzipFile(None, gz_mode, compresslevel, filename)

File c:\Python312\Lib\gzip.py:192, in GzipFile.__init__(self, filename, mode, compresslevel, fileobj, mtime)
    190     mode += 'b'
    191 if fileobj is None:
--> 192     fileobj = self.myfileobj = builtins.open(filename, mode or 'rb')
    193 if filename is None:
    194     filename = getattr(fileobj, 'name', '')

FileNotFoundError: [Errno 2] No such file or directory: 'models\\1c.pklz'
```

## 2 Prepositional logic (2đ)

In this session, you will practice transforming a natural languages into a first-order logic formula. For each successful assertion, you will receive 0.5 points.

For example: "There is a lighting source that shines"

```
def lightShines():
    def Light(x): return Atom('Light', x)    # 
    def Shines(x): return Atom('Shines', x)  # 
    return Exists('$x', And(Light('$x'), Shines('$x')))
```

##### **2a) Everybody has mother**

<!-- grade_id: TODO_2a, points: 0, grade: false, solution: true -->
```python
def formula2a() -> Formula:
    # Predicates to use:
    def Person(x): return Atom('Person', x)        # whether x is a person
    def Mother(x, y): return Atom('Mother', x, y)  # whether x's mother is y

    # Note: You do NOT have to enforce that the mother is a "person"
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: GRADE_2a, points: 0.5, grade: true, solution: false -->
```python
formula2a_precondition = AntiReflexive('Mother')
assert checkFormula('2a', formula2a(), formula2a_precondition) == True
```

##### **2b) At least one person does not have any child**

<!-- grade_id: TODO_2b, points: 0, grade: false, solution: true -->
```python
def formula2b() -> Formula:
    # Predicates to use:
    def Person(x): return Atom('Person', x)        # whether x is a person
    def Child(x, y): return Atom('Child', x, y)    # whether x has a child y

    # Note: You do NOT have to enforce that the child is a "person"
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: GRADE_2b, points: 0.5, grade: true, solution: false -->
```python
formula2b_precondition = AntiReflexive('Child')
assert checkFormula('2b', formula2b(), formula2b_precondition) == True
```

##### **2c) Create definition for Daughter by Female and Child**

Examples

```
# Defnition of Parent by Child
def parentChild():
    def Parent(x, y): return Atom('Parent', x, y)  # whether x has a parent y
    def Child(x, y): return Atom('Child', x, y)    # whether x has a child y
    return Forall('$x', Forall('$y', Equiv(Parent('$x', '$y'), Child('$y', '$x'))))
    
```

<!-- grade_id: TODO_2c, points: 0, grade: false, solution: true -->
```python
def formula2c() -> Formula:
    # Predicates to use:
    def Female(x): return Atom('Female', x)            # whether x is female
    def Child(x, y): return Atom('Child', x, y)        # whether x has a child y
    def Daughter(x, y): return Atom('Daughter', x, y)  # whether x has a daughter y
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: GRADE_2c, points: 0.5, grade: true, solution: false -->
```python
formula2c_precondition = AntiReflexive('Child')
assert checkFormula('2c', formula2c(), formula2c_precondition) == True
```

##### **2d) Create definition for Grandmother by Female and Parent**

<!-- grade_id: TODO_2d, points: 0, grade: false, solution: true -->
```python
# Note: It is ok for a person to be her own parent
def formula2d() -> Formula:
    # Predicates to use:
    def Female(x): return Atom('Female', x)                  # whether x is female
    def Parent(x, y): return Atom('Parent', x, y)            # whether x has a parent y
    def Grandmother(x, y): return Atom('Grandmother', x, y)  # whether x has a grandmother y
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: GRADE_2d, points: 0.5, grade: true, solution: false -->
```python
formula2d_precondition = AntiReflexive('Parent')
assert checkFormula('2d', formula2d(), formula2d_precondition) == True
```

## 3. Who is the liar? (2.5đ)

Using prepositional method for the following problem:

A Facebook employee make a serious mistake that crashed the server, people are discussing to find out the one that caused this problem. After a while, four suspects are identified.

We have the following information:
- John: "I am not the guy that make the server down!"
- Susan: "It should be Nicole."
- Mark: "No, it was Susan!"
- Nicole: "Susan is lying, you should not trust her."
- There is **only 1 person** told the truth.
- There is **only 1 person** crashed the server.

Finish the ```liar()``` function with 6 formulas for finding the liar!

<!-- grade_id: TODO_3, points: 0, grade: false, solution: true -->
```python
# Problem 3: Liar puzzle

# Facts:
# 0. John: "It wasn't me!"
# 1. Susan: "It was Nicole!"
# 2. Mark: "No, it was Susan!"
# 3. Nicole: "Susan's a liar."
# 4. Exactly one person is telling the truth.
# 5. Exactly one person crashed the server.
# Query: Who did it?

# This function returns a list of 6 formulas corresponding to each of the
# above facts. Be sure your formulas are exactly in the order specified. 


# Hint: For fact 4 & 5, you might want to use the Equals predicate, defined as:
# def Equals(x, y): return Atom('Equals', x, y)
# This predicate is used to assert that two objects are the same.
# In particular, Equals(x,x) = True and Equals(x,y) = False iff x is not equal to y. 
# It can also be solved in other ways, without the Equals predicate!
    
def liar() -> Tuple[List[Formula], Formula]:
    def TellTruth(x): return Atom('TellTruth', x)
    def CrashedServer(x): return Atom('CrashedServer', x)
    john = Constant('john')
    susan = Constant('susan')
    nicole = Constant('nicole')
    mark = Constant('mark')
    formulas = []
    # We provide the formula for fact 0 here.
    formulas.append(Equiv(TellTruth(john), Not(CrashedServer(john))))
    
    # You should add 5 formulas, one for each of facts 1-5.
    
    # YOUR CODE HERE
    raise NotImplementedError()
    query = CrashedServer('$x')
    return (formulas, query)
```

<!-- grade_id: GRADE_3_len, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
assert len(predForms) == 6
```

```python
predForms, predQuery = liar()
formula_id = 0
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

<!-- grade_id: GRADE_3_1, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 1
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

<!-- grade_id: GRADE_3_2, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 2
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

<!-- grade_id: GRADE_3_3, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 3
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

<!-- grade_id: GRADE_3_4, points: 0.5, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 4
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

<!-- grade_id: GRADE_3_5, points: 0.5, grade: true, solution: false -->
```python
predForms, predQuery = liar()
formula_id = 5
assert checkFormula('3a-' + str(formula_id), predForms[formula_id]) == True
```

<!-- grade_id: GRADE_3_ALL, points: 0.25, grade: true, solution: false -->
```python
predForms, predQuery = liar()
assert checkFormula('3a-all', AndList(predForms)) == True
```

```python
# Run this cell to solve the puzzle and find out who was the liar!
predForms, predQuery = liar()
kb = createModelCheckingKB()

filename = os.path.join('models', '3a-all.pklz')
objects, targetModels = pickle.load(gzip.open(filename))
for obj in objects:
    kb.tell(Atom('Object', obj))

for predForm in predForms:
    response = kb.tell(predForm)

response = kb.ask(predQuery)
showKBResponse(response)
```

## 4. Inference (4.75đ)

**You will designed this part by your own python scripts. Please do not write to this notebook**

Given a Knowledge Base **(KB)** and a statement $\alpha$, both represented by prepositional logic and change to CNF. Determine if KB entails $\alpha$ (KB ⊨ $\alpha$) by resolution.

- I have provided you with seven test cases, of which I have included examples of expected outputs for the first two test cases (input01.txt, input02.txt). Your task is to determine the outputs for all test cases and present them in two ways: 
    + Write the output explanation in your report. 
    + Design your own Python scripts to generate these outputs from given inputs.   

- Write short evaluations about the advantages and disadvantages of resolution method for prepositional logic, propose your own solution for specific problem.

**a. Desribe the experiment data.**
- **Input data**: **KB** and $\alpha$ with CNF stored in **input.txt**. This file have the following structure:
    - The first line contains $\alpha$ statement
    - The second line has N – the number of clauses in the KB.
    - N following lines are clauses in KB, one line for each clause.
    
Positive literal is represented by an uppercase character (A-Z). Negative literal is positive literal with minus (‘-‘) before the character.

Literals are linked by OR. Bettween literals and keywords there can be more than one space.

- **Output data**: The set of statements that are generated during the whole resolution process and the conclusion clause is stored in **output.txt**. This file must have the following structure:
    - The first line contains M1 – the number of generated clauses in first loop. M1 following lines describe those clauses (including empty clause), each line represents a clause. The empty clause is represented by “{}”
    - Next loops are also represented as ($M_2$, $M_3$,…, $M_n$ clauses) the first loop. 
    - The last line is the conclusion clause, which mean if “KB entails $\alpha$?”. Print YES if KB entails $\alpha$. NO for the opposite.
    - Remove redundant clause (same clauses in same loop, or the loops before).
    
    
- **Main function must**:
    - Read the input data and store it in a suitable data structure.
    - Call PL-Resolution function for execute resolution algorithm.
    - Export the output data with the required format.

- **Other important notes**:
    - Store the semantic information of true and false value in PL-RESOLUTION. Do not forget that we need to negate $\alpha$.
    - Literals in a clause is sorted by alphabetical ordering.
    - Infered condition is checked in each loop as when all the new clauses are generated from the current KB.
    - Clauses that have A V B V -B format and True value are similar to A V True can be removed.


**b. Evaluations**
- Read input data and store in suitable data structure, you will recieve 0.5
- Implementation of resolution method, you will recieve 0.75
- Each test case with a complete and correct explanation of the output will receive 0.5.
- If the report does not include a detailed explanation of how the output was obtained, that test case will not be scored.
- **If you only include the output explanation in the report without the code, you will receive 0.25 for each test case.**

**<font color='red'>In cases where submissions are identical or cheating is detected, you will receive 0 points for the entire course. Therefore, remember to reference any sources you consult and do not cheat!</font>**